Vglut2-based glutamatergic signaling in central noradrenergic neurons is dispensable for normal breathing and chemosensory reflexes

Author response:

The following is the authors’ response to the previous reviews.

All of the reviewers indicate that their major concerns have been adequately addressed, but they each have a few comments that the authors should consider before submitting a final version (without further review) for publication. For example, a statement about the sex of the mice used in the studies and whether any differences were noted if both sexes were used. The idea that the loss of glutamate transport might affect NA loading into vesicles is also worth considering. Finally, the authors might want to mention that the role of neuropeptide release from NA neurons needs further examination.

As noted in the prior submitted revision, all experiments contained both males and females and this was addressed in our re-submission. In our analysis of breathing and metabolism, sex was included in the analysis and no significant phenotypic difference was observed (The statement of no sex difference is in line 451-456). For the fate map and in situ experiments, although the group size is small, we did not see obvious differences in the expression patterns in the three glutamate transporters between females and males (line 347-350). All the anatomical and phenotypic data in this manuscript are presented as combined graphs (figure 1, figure 1 supplement 1, figure 2, figure 2 supplement 2, figure 4,5,6,7) and we had differentially labeled our data points by sex (female data is pink and male data is blue).

The possibility that loss of Vglut2 might affect NA release has been added in the discussion (line 485-491) of the current revision. Dopamine Beta Hydroxylase (DBH) converts dopamine to noradrenaline in the vesicles, thus, glutamate may not directly affect noradrenaline loading into vesicles. However, since loss of Vglut2 reduced dopamine release in subsets of dopaminergic neurons, it remains possible that glutamate affects dopamine loading in NA neurons and in turn perturbs DA to NA conversion in the vesicle by DBH and subsequent noradrenaline release. Future work could examine this hypothesis using fast-scan cyclic voltammetry (FSCV) or microdialysis.

The further examination of the role of neuropeptide release from NA neurons is mentioned in the discussion (line 491-494 and line 497-499 of the pre).

eLife assessment

Chang et al. provide glutamate co-expression profiles in the central noradrenergic system and test the requirement of Vglut2-based glutamatergic release in respiratory and metabolic activity under physiologically relevant gas challenges. Their experiments provide compelling evidence that conditional deletion of vesicular glutamate transporters from noradrenergic neurons does not impact steady-state breathing or metabolic activity in room air, hypercapnia, or hypoxia. This study provides an important contribution to our understanding of how noradrenergic neurons regulate respiratory homeostasis in conscious adult mice.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

Chang et al. provide glutamate co-expression profiles in the central noradrenergic system and test the requirement of Vglut2-based glutamatergic release in respiratory and metabolic activity under physiologically relevant gas challenges. Their experiments show that conditional deletion of Vglut2 in NA neurons does not impact steady-state breathing or metabolic activity in room air, hypercapnia, or hypoxia. Their observations challenge the importance of glutamatergic signaling from Vglut2 expressing NA neurons in normal respiratory homeostasis in conscious adult mice.

Strengths:

The comprehensive Vglut1, Vglut2, and Vglut3 co-expression profiles in the central noradrenergic system and the combined measurements of breathing and oxygen consumption are two major strengths of this study. Observations from these experiments provide previously undescribed insights into (1) expression patterns for subtypes of the vesicular glutamate transporter protein in the noradrenergic system and (2) the dispensable nature of Vglut2dependent glutamate signaling from noradrenergic neurons to breathing responses to physiologically relevant gas challenges in adult conscious mice.

Weaknesses:

Although the cellular expression profiles for the vesicular glutamate transporters are provided, the study does not document that glutamatergic-based signaling originating from noradrenergic neurons is evident at the cellular level under normal, hypoxic, and/or hypercapnic conditions. The authors effectively recognize this issue and appropriately discuss their findings in this context.

We thank the reviewer for the positive evaluation of our work.

Reviewer #2 (Public Review):

The authors characterized the recombinase-based cumulative fate maps for vesicular glutamate transporters (Vglut1, Vglut2 and Vglut3) expression and compared those maps to their realtime expression profiles in central NA neurons by RNA in situ hybridization in adult mice. Authors have revealed a new and intriguing expression pattern for Vglut2, along with an entirely uncharted co-expression domain for Vglut3 within central noradrenergic neurons. Interestingly, and in contrast to previous studies, the authors demonstrated that glutamatergic signaling in central noradrenergic neurons does not exert any influence on breathing and metabolic control either under normoxic/normocapnic conditions or after chemoreflex stimulation. Also, they showed for the first-time the Vglut3-expressing NA population in C2/A2 nuclei. In addition, they were also able to demonstrate Vglut2 expression in anterior NA populations, such as LC neurons, by using more refined techniques, unlike previous studies.

A major strength of the study is the use of a set of techniques to investigate the participation of NA-based glutamatergic signaling in breathing and metabolic control. The authors provided a full characterization of the recombinase-based cumulative fate maps for Vglut transporters. They performed real-time mRNA expression of Vglut transporters in central NA neurons of adult mice. Further, they evaluated the effect of knocking down Vglut2 expression in NA neurons using a DBH-Cre; Vglut2cKO mice on breathing and control in unanesthetized mice. Finally, they injected the AAV virus containing Cre-dependent Td tomato into LC of v-Glut2 Cre mice to verify the VGlut2 expression in LC-NA neurons. A very positive aspect of the article is that the authors combined ventilation with metabolic measurements. This integration holds

particular significance, especially when delving into the exploration of respiratory chemosensitivity. Furthermore, the sample size of the experiments is excellent. Despite the clear strengths of the paper, some weaknesses exist. It is not clear in the manuscript if the experiments were performed in males and females and if the data were combined. I believe that the study would have benefited from a more comprehensive analysis exploring the sex specific differences. The reason I think this is particularly relevant is the developmental disorders mentioned by the authors, such as SIDS and Rett syndrome, which could potentially arise from disruptions in central noradrenergic (NA) function, exhibit varying degrees of sex predominance. Moreover, some of the noradrenergic cell groups are sexually dimorphic. For instance, female Wistar rats exhibit a larger LC size and more LC-NA neurons than male subjects (Pinos et al., 2001; Garcia-Falgueras et al., 2005). More recently, a detailed transcriptional profiling investigation has unveiled the identities of over 3,000 genes in the LC. This revelation has highlighted significant sexual dimorphisms, with more than 100 genes exhibiting differential expression within LC-NA neurons at the transcript level. Furthermore, this investigation has convincingly showcased that these distinct gene expression patterns have the capacity to elicit disparate behavioral responses between sexes (Mulvey et al., 2018).

Therefore, the authors should compare the fate maps, Vglut transporters in males and females, at least considering LC-NA neurons. Even in the absence of identified sex differences, this information retains significant importance.

An important point well raised by the authors is that although suggestive, these experiments do not definitively rule out that NA-Vglut2 based glutamatergic signaling has a role in breathing control. Subsequent experiments will be necessary to validate this hypothesis.

An improvement could be made in terms of measuring body temperature. Opting for implanted sensors over rectal probes would circumvent the need to open the chamber, thereby preventing alterations in gas composition during respiratory measurements. Further, what happens to body temperature phenotype in these animals under different gas exposures? These data should be included in the Tables.

Is it plausible that another neurotransmitter within NA neurons might be released in higher amounts in DBH-Cre; Vglut2 cKO mice to compensate for the deficiency in glutamate and prevent changes in ventilation?

Continuing along the same line of inquiry is there a possibility that Vglut2 cKO from NA neurons not only eliminates glutamate release but also reduces NA release? A similar mechanism was previously found in VGLUT2 cKO from DA neurons in previous studies (Alsio et al., 2011; Fortin et al., 2012; Hnasko et al., 2010). Additionally, does glutamate play a role in the vesicular loading of NA? Therefore, could the lack of effect on breathing be explained by the lack of noradrenaline and not glutamate?

We thank the reviewer for the positive evaluation and further suggestions. Please see our response in “Author Response” to the previous version of Reviewer #2 (Public review).

Reviewer #4 (Public Review):

Summary:

Although previous research suggested that noradrenergic glutamatergic signaling could influence respiratory control, the work performed by Chang and colleagues reveals that excitatory (specifically Vglut2) neurons is dynamically and widely expressed throughout the central noradrenergic system, but it is not significantly crucial to change baseline breathing as well the hypercapnia and hypoxia ventilatory responses. The central point that will make a significant change in the field is how NA-glutamate transmission may influence breathing control and the dysfunction of NA neurons in respiratory disorders.

Strengths:

There are several strengths such as the comprehensive analysis of Vglut1, Vglut2, and Vglut3 expression in the central noradrenergic system and the combined measurements of breathing parameters in conscious unrestrained mice.

Other considerations :

These results strongly suggest that glutamate may not be necessary for modulating breathing under normal conditions or even when faced with high levels of carbon dioxide (hypercapnia) or low oxygen levels (hypoxia). This finding is unexpected, considering many studies have underscored glutamate's vital role in respiratory regulation, more so than catecholamines. This leads us to question the significance of catecholamines in controlling respiration. Moreover, if glutamate is not essential for this function, we need to explore its role in other physiological processes such as sympathetic nerve activity (SNA), thermoregulation, and sensory physiology.

We thank the reviewer for the positive evaluation and further suggestions. The potential role of noradrenergic-derived glutamate in other processes, which is beyond the scope of this study, should be addressed in the future.

Recommendations for the authors:

Reviewer #2 (Recommendations For The Authors):

All of my concerns were effectively resolved, leading me to accept the paper. However, I suggest that the authors consider investing in a more reliable system for measuring body temperature, as accurate measurements of this parameter are crucial for whole body plethysmography.

Thank you for the suggestion. The real-time measurement of body temperature is a goal in future studies.

Reviewer #4 (Recommendations For The Authors):

Because I am revising a revised version, I believe the authors have addressed most, if not all, the concerns raised by already 3 reviewers. In my understanding the authors achieved their aims and the results are totally supported by the conclusions. The impact of this work on the respiratory field is significant and is likely to advance the field. The methods and data utilized, which combine standard techniques with genetic tools, will be highly beneficial to the research community.

In my understanding I still have one concern that if glutamate is not critical, then what is? Could we potentially disable the noradrenergic (NA) system while preserving glutamate functionality to determine if the NA system is indeed crucial for respiratory physiology? This approach might provide clearer insights into the mechanisms underlying respiratory control.

We agree that there remain several exciting questions about the respective roles of noradrenaline, glutamate, and other neuropeptides such as Neuropeptide Y (NPY) and galanin. We are currently devising strategies to address the respective and combinatorial roles for all these candidates in breathing control. Most simply, we can conditionally, mutagenized each of them in the central noradrenergic system in an acute manner using DBH-CreER mice to determine if any of them are critical to respiratory control with the advantage of minimizing developmental compensatory events.

Changes from the earlier version to the revised version

Abstract

Abstract

Central noradrenergic (NA) neurons are key constituents of the respiratory homeostatic network. NA dysfunction is implicated in several developmental respiratory disorders including Congenital Central Hyperventilation Syndrome (CCHS), Sudden Infant Death Syndrome (SIDS) and Rett Syndrome. The current unchallenged paradigm in the field, supported by multiple studies, is that glutamate co-transmission in subsets of central NA neurons plays a role in breathing control. If true, NA-glutamate co-transmission may also be mechanistically important in respiratory disorders. However, the requirement of NA-derived glutamate in breathing has not been directly tested and the extent of glutamate co-transmission in the central NA system remains uncharacterized. Therefore, we fully characterized the cumulative fate maps and acute adult expression patterns of all three Vesicular Glutamate Transporters (Slc17a7 (Vglut1), Slc17a6 (Vglut2), and Slc17a8 (Vglut3)) in NA neurons, identifying a novel, dynamic expression pattern for Vglut2 and an undescribed co-expression domain for Vglut3 in the NA system. In contrast to our initial hypothesis that NA derived glutamate is required to breathing, our functional studies showed that loss of Vglut2 throughout the NA system failed to alter breathing or metabolism under room air, hypercapnia, or hypoxia in unrestrained and unanesthetized mice. These data demonstrate that Vglut2-based glutamatergic signaling within the central NA system is not required for normal baseline breathing and hypercapnic, hypoxic chemosensory reflexes. These outcomes challenge the current understanding of central NA neurons in the control of breathing and suggests that glutamate may not be a critical target to understand NA neuron dysfunction in respiratory diseases.

Introduction

Breathing is a vital and life-sustaining function supporting homeostatic processes, most critically maintaining blood pH/CO2 and O2 levels within a narrow physiological range. Respiratory homeostasis is mediated by neuron-modulated lung ventilation adjustments in response to physiological deviations resulting in high pCO2 or low pO2 blood and tissue levels, known as the hypercapnic and hypoxic reflexes respectively (Del Negro et al., 2018; Dick et al., 2018). These reflexes are part of a complex brainstem neural network that integrates a multitude of information streams across the central and peripheral nervous systems to regulate respiratory output. Within this brainstem network, central noradrenergic (NA) neurons are known to be an important component that plays a variety of roles in modulating breathing. Furthermore, various perturbations across the central NA system have been implicated in several developmental disorders with respiratory and chemosensory features such as Congenital Central Hyperventilation Syndrome (CCHS), Sudden Infant Death Syndrome (SIDS) and Rett Syndrome (Beltrán-Castillo et al., 2017; Feldman et al., 2013; Gauda et al., 2007; Viemari, 2008). Thus, understanding how central NA neurons modulate respiratory chemoreflexes or chemosensory breathing is critically important for the development of new diagnostic and therapeutic interventions to address respiratory pathophysiology.

Central noradrenergic neurons are commonly thought to exert their effect on breathing through their primary neurotransmitter noradrenaline and adrenaline. However, in addition to noradrenaline, several studies provide strong evidence that glutamate is co-transmitted in subsets of central NA neurons. Vesicular glutamate transporter 2 (Vglut2), a gene marker of glutamatergic signaling, has been shown to be co-expressed in subsets of central NA neurons including, C1/A1, C2/A2, A5, and LC in adult rats and mice (DePuy et al., 2013; Souza et al., 2022a; Stornetta et al., 2002a, 2002b; Yang et al., 2021), though NA specific expression of related Vglut1 and Vglut3 transporters remains unknown. Additionally, it has been well documented that central NA neurons co-expressing Vglut2 innervate key respiratory centers, such as preBötzinger complex and parafacial region (pFRG). Additionally, central NA neurons co-expressing Vglut2 project to many key autonomic brainstem, spinal cord, and forebrain centers, such as the dorsal motor nucleus of the vagus, intermediolateral nucleus (IML) and sympathetic system, and the hypothalamus, all of which could drive a change in breathing if perturbed, i.e., indirect effects from cardio-vascular and metabolic dysregulation (Supplemental table 1). Thus, it has become a dominant paradigm in the field of respiratory physiology that NA-based glutamate release is a critical form of neurotransmission in the control of breathing. Abbott et al. (2014) showed that Vglut2 is required for an increase in respiratory rate when anterior C1 neurons are unilaterally opto-genetically stimulated. Malheiros-Lima et al. (2020) showed that Vglut2-expressing C1 neurons project to the pFRG region, and hypoxic breathing was blunted after blockade of ionotropic glutamatergic receptors at the pFRG site in anesthetized rats, together supporting a role for anterior C1 neurons releasing glutamate at the pFRG site to regulate breathing under hypoxia. Similarly, Malheiros-Lima et al. (2022) and Malheiros-Lima et al. (2018) showed that Vglut2-expressing C1 neurons project to the NA A5 region and the preBötzinger complex, and, again, the blockade of ionotropic glutamatergic receptors at the A5 region or preBötzinger complex reduced the increase in phrenic nerve activity and respiratory frequency elicited by optogenetic stimulation of C1 cells in an anesthetized preparation. In addition, Guyenet et al. (2013) and others speculated that the apparent lack of plasmalemmal monoamine transporter in C1 fibers indicates reduced or absent noradrenergic or adrenergic signaling due a lack of re-uptake and neurotransmitter pool depletion (Comer et al., 1998; Lorang et al., 1994). Cumulatively, the studies argue for glutamate as the predominant functional neurotransmitter for C1 NA neurons in the breathing neural network. To our knowledge, this dominant perspective has not been otherwise previously challenged. Although these studies are informative, the evidence supporting the role of NA-based Vglut2 signaling in respiratory control are either indirect and circumstantial or cannot be seen as physiological given experimental limitations, such as the focal nature of opto-genetic stimulation. Thus, it is not yet clear what the requirement is for NA-based glutamatergic signaling in homeostatic breathing in the unanesthetized and unrestrained animal and how that might inform upon disease.

To better understand the role of NA-based glutamatergic signaling in breathing, we sought to both fully characterize the molecular profiles of the central noradrenergic system with respect to glutamate co-expression and to test the hypothesis that Vglut2-based glutamatergic release is required in respiratory control under physiological chemosensory challenges in unanesthetized and unrestrained mice. We first fully characterized the recombinase-based cumulative fate maps for Vglut1, Vglut2 and Vglut3 expression and compared those maps to their real time expression profiles in central NA neurons by RNA in situ hybridization in adult mice. We found a novel dynamic expression pattern for Vglut2 and an entirely undescribed co-expression domain for Vglut3 in the central NA system. Secondly, to determine if Vglut2-based glutamatergic signaling in NA neurons is required for respiratory homeostasis, we conditionally ablated Vglut2 in all NA neurons and tested respiratory, chemosensory, and metabolic function in unrestrained and unanesthetized mice. Using the same genetic model in prior studies (Abbott et al., 2014), conditional deletion of Vglut2 in NA neurons did not significantly impact breathing under room air, hypercapnic, or hypoxic conditions. These results demonstrate, for the first time, that NA Vglut2-based glutamatergic signaling is dispensable for respiratory control, which challenges the prevalent perspective on the role of C1 and other Vglut2 expressing NA neurons in respiratory homeostasis and suggests that glutamate may not be a critical target to understand NA neuron dysfunction in respiratory diseases.

Results

Cumulative fate maps of central noradrenergic neurons co-expressing Vglut1, Vglut2 and Vglut3

Added in the revised version: To fully characterize the expression profiles of all three glutamate markers Vglut1, Vglut2 and Vglut3 in central noradrenergic neurons, we first used an intersectional genetic strategy (Figure 1A). We bred three Cre drivers Slc17a7Cre(Vglut1-Cre), Slc17a6Cre (Vglut2-Cre) and Slc17a8Cre(Vglut3-Cre) to Dbhp2a-Flpo (DBH-p2a-Flpo) (targeting noradrenergic neurons) mice, respectively. The three compound lines of Slc17a7Cre; Dbhp2a-Flpo, Slc17a6Cre; Dbhp2a-Flpo, and Slc17a8Cre; Dbhp2a- Flpo were then crossed with the Rosa26RC::FLTG (RC::FLTG). RC::FLTG are intersectional reporter mice that express tdTomato in cells expressing only flippase (Flpo) and express eGFP in cells co-expressing both Flpo and Cre recombinases. Thus, in each of the three intersectional reporter crosses [e.g. 1) Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, 2) Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG and 3) Slc17a8Cre; Dbhp2a-Flpo; Rosa26RC::FLTG], NA neurons without any Vglut1, 2, or 3 expression are labelled by red fluorescent protein (tdTomato) while NA neurons co-expressing either Vglut1, Vglut2, or Vglut3 are labeled by green fluorescent protein (eGFP). By using this method, we characterized and quantified the expression profiles of all three vesicular glutamate transporters Vglut1, Vglut2, and Vglut3 across every anatomically defined NA nucleus in adult mice including A7, Locus Coeruleus (LC), dorsal/ ventral subcoeruleus nucleus (sub CD/CV), A5, and the anterior - posterior dimensions of C1/A1, C2/A2 (Robertson et al., 2013), and C3 (Figure 1B-D and Figure 1 supplement 1A-B). Vglut1-Cre was not co-expressed in any central NA nuclei in adult mice. However, Vglut2-Cre and Vglut3-Cre both showed co-expression in central NA neurons. Vglut3-Cre-expressing NA neurons were restricted to posterior C2/A2 and posterior C1/A1, with greatest expression in posterior C2/A2 where 26.9 ± 3.16% of NA neurons were Vglut3-Cre positive. In posterior C1/A1, only 1.26 ± 0.559% of NA neurons were Vglut3-Cre positive. Surprisingly, 84.6 ± 3.75% of NA neurons in total showed Vglut2-Cre co-expression and each NA nucleus was predominantly labelled by Vglut2-Cre expression. Over 50% of NA neurons in each NA nucleus were Vglut2-Cre positive. The percentages of Vglut2-Cre positive NA neurons in each NA nucleus were as follows: 69.8 ± 8.29% in A7, 80.5 ± 3.78% in LC, 84.8 ± 11.3% in A5, 51.5 ± 2.61% in sub CD/CV, 99.0 ± 0.681% in anterior C1/A1, 95.7 ± 2.40% in posterior C1/A1, 96.1 ± 1.98% in anterior C2/A2, 97.8 ± 0.771% in posterior C2/A2, and 100 ± 0.00% in C3 (Mean ± SEM). The presence of Vglut2-Cre co-expression in anterior NA groups was unexpected as previous in situ data in adult rats found Slc17a6 (Vglut2) positive NA neurons only in the posterior C2/A2 and C1/A1 (Stornetta et al., 2002a, 2002b). Notably, however, the Vglut2 co-expression in the LC region agrees with Yang et al. (2021) which also showed that 89% of LC NA neurons are Vglut2 positive using a different intersectional strategy in mice: ThFlpo; Slc17a6Cre; Rosa26Ai65 (TH-Flpo; Vglut2-Cre; Ai65). The differences in our and other fate maps compared to in situ hybridization may either reflect early gene expression that is down regulated in the adult or may reflect low levels of expression not detectable by in situ hybridization but that are nonetheless sufficient to affect recombination in the intersectional genetic strategy.

Removed in the revised version: To fully characterize the expression profiles of all three glutamate markers Vglut1, Vglut2 and Vglut3 in central noradrenergic neurons, we first used an intersectional genetic strategy (Figure 1A). We bred three Cre drivers Slc17a7Cre(Vglut1-Cre), Slc17a6Cre (Vglut2-Cre) and Slc17a8Cre(Vglut3-Cre) to Dbhp2a-Flpo (DBH-p2a-Flpo) (targeting noradrenergic neurons) mice, respectively.

Cumulative fate maps of glutamate co-expressing NA neurons characterized by intersectional genetics.

(A) Breeding schematic to generate the three intersectional reporter lines Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a8Cre; Dbhp2a- Flpo; Rosa26RC::FLTG. In each of the three intersectional reporter crosses, NA neurons co-expressing either Vglut1, Vglut2, or Vglut3 are labeled by green fluorescent protein (eGFP) while NA neurons without any Vglut1,2, or 3 expressions are labelled by red fluorescent protein (tdTomato). Shown as an example is the Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG (Vglut2-Cre/+; DBH-p2a-Flpo/+; RC::FLTG/+) intersectional reporter line. (B) Quantification of the percentage of Vglut2-coexpressing NA neurons among NA neurons in each brainstem NA nuclei including A7, LC, A5, sub CD/CV, anterior C1/A1 and C2/A2, posterior C1/A1 and C2/A2. Pink data points represent female data while blue data points represent male data. (C) Quantification of the percentage of Vglut3-coexpressing NA neurons among NA neurons in each NA nuclei. Pink data points represent female data while blue data points represent male data. (D) Fluorescent expression of tdTomato (red) and eGFP (green) in coronal sections of Vglut1, Vglut2, or Vglut3 intersectional reporter lines in brainstem NA nuclei. Nucleus is labeled by DAPI (blue). Scale bar 50μm.

Removed in the revised version: The three compound lines of Slc17a7Cre; Dbhp2a-Flpo, Slc17a6Cre; Dbhp2a-Flpo, and Slc17a8Cre; Dbhp2a- Flpo were then crossed with the Rosa26RC::FLTG (RC::FLTG). RC::FLTG are intersectional reporter mice that express tdTomato in cells expressing only flippase (Flpo) and express eGFP in cells co-expressing both Flpo and Cre recombinases. Thus, in each of the three intersectional reporter crosses [e.g. 1) Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, 2) Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG and 3) Slc17a8Cre; Dbhp2a-Flpo; Rosa26RC::FLTG], NA neurons without any Vglut1, 2, or 3 expression are labelled by red fluorescent protein (tdTomato) while NA neurons co-expressing either Vglut1, Vglut2, or Vglut3 are labeled by green fluorescent protein (eGFP). By using this method, we characterized and quantified the expression profiles of all three vesicular glutamate transporters Vglut1, Vglut2, and Vglut3 across every anatomically defined NA nucleus in adult mice including A7, Locus Coeruleus (LC), dorsal/ ventral subcoeruleus nucleus (sub CD/CV), A5, and the anterior - posterior dimensions of C1/A1, C2/A2 (Robertson et al., 2013), and C3 (Figure 1B-D and Figure 1 supplement 1A-B). Vglut1-Cre was not co-expressed in any central NA nuclei in adult mice. However, Vglut2-Cre and Vglut3-Cre both showed co-expression in central NA neurons. Vglut3-Cre-expressing NA neurons were restricted to posterior C2/A2 and posterior C1/A1, with greatest expression in posterior C2/A2 where 26.9 ± 3.16% of NA neurons were Vglut3-Cre positive. In posterior C1/A1, only 1.26 ± 0.559% of NA neurons were Vglut3-Cre positive. Surprisingly, 84.6 ± 3.75% of NA neurons in total showed Vglut2-Cre co-expression and each NA nucleus was predominantly labelled by Vglut2-Cre expression. Over 50% of NA neurons in each NA nucleus were Vglut2-Cre positive. The percentages of Vglut2-Cre positive NA neurons in each NA nucleus were as follows: 69.8 ± 8.29% in A7, 80.5 ± 3.78% in LC, 84.8 ± 11.3% in A5, 51.5 ± 2.61% in sub CD/CV, 99.0 ± 0.681% in anterior C1/A1, 95.7 ± 2.40% in posterior C1/A1, 96.1 ± 1.98% in anterior C2/A2, 97.8 ± 0.771% in posterior C2/A2, and 100 ± 0.00% in C3 (Mean ± SEM). The presence of Vglut2-Cre co-expression in anterior NA groups was unexpected as previous in situ data in adult rats found Slc17a6 (Vglut2) positive NA neurons only in the posterior C2/A2 and C1/A1 (Stornetta et al., 2002a, 2002b). Notably, however, the Vglut2 co-expression in the LC region agrees with Yang et al. (2021) which also showed that 89% of LC NA neurons are Vglut2 positive using a different intersectional strategy in mice: ThFlpo; Slc17a6Cre; Rosa26Ai65 (TH-Flpo; Vglut2-Cre; Ai65). The differences in our and other fate maps compared to in situ hybridization may either reflect early gene expression that is down regulated in the adult or may reflect low levels of expression not detectable by in situ hybridization but that are nonetheless sufficient to affect recombination in the intersectional genetic strategy.

Real time mRNA expression patterns of Vglut1, Vglut2 and Vglut3 in central noradrenergic neurons in adult mice

To verify if the real time NA-based Vglut2 expression pattern in adult mice is comparable to previous reports in adult rats (DePuy et al., 2013; Stornetta et al., 2002b, 2002a), and to further characterize the expression patterns of the other two vesicular glutamate transporters (Vglut1 and Vglut3) in adult mice, we performed fluorescent RNA in situ hybridization experiments. We co-stained Slc17a7 (Vglut1), Slc17a6 (Vglut2), or Slc17a8 (Vglut3) with Dbh respectively in brain tissue of adult mice and characterized the colocalization of Slc17a7, Slc17a6, Slc17a8 with Dbh in brainstem NA nuclei A7, LC, A5, sub CD/CV, anterior C1/A1 and C2/A2, posterior C1/A1 and C2/A2, and C3 (Figure 2A-C and Figure 2 supplement 2A-B). No Slc17a7/Dbh double positive neurons were detected in any part of the central NA system, consistent with our fate map. Slc17a8 mRNA colocalization with Dbh was only found in the posterior part of C2/A2 where 27.1 ± 1.86% of Dbh positive neurons demonstrated Slc17a8 colocalization, again consistent with our fate map. However, we only found detectable levels of Slc17a6 mRNA in the C1/A1, C2/A2 and C3 NA regions, but not in the A7, LC, A5, sub CD/CV. The percentages of Slc17a6/Dbh double positive neurons for regions where detected were as follows: 84.7 ± 5.82% in anterior C1/A1, 66.3 ± 0.335% in anterior C2/A2, 35.7 ± 3.01% in posterior C1/A1, 90.1 ± 2.45% in posterior C2/A2, and 79.7 ± 3.94% in C3. This result is consistent with the previous in situ data in adult rats (Stornetta et al., 2002a, 2002b) suggesting that there is no obvious expression difference of Slc17a6 in the central NA system between mouse and rat. However, the Slc17a6 in situ data did not show expression in anterior NA populations. This difference between our in situ data and fate map data supports our previous hypothesis that many NA neurons expressed Slc17a6 (Vglut2) at some point from early development toward adulthood but the expression of Slc17a6 (Vglut2) is diminished during adulthood. Interestingly, Yang et al. (2021) did show adult Vglut2 co-expression in LC by a viral injection of both Cre and Flpo-dependent eYFP into the LC of a bi-transgenic mouse with both ThFlpoand Slc17a6Cre. To verify this result, we injected an AAV virus containing Cre-dependent tdTomato (pAAV-EF1a-DIO-tdTomato-WPRE) into the LC region of adult Slc17a6Cre mice. Consistent with Yang et al. (2021), we observed sparse Vglut2-Cre and TH-immunopositive double labelled neurons in the LC (Figure 2 supplement 1). Additionally, Souza et al. (2022a) showed about 28.5% of A5 neurons are Slc17a6 mRNA positive in adult rats by RNA scope, a method which is more sensitive than the traditional in situ hybridization we used here. These data suggest that some adult NA neurons in LC and A5 have Slc17a6 (Vglut2) co-expression, but the Slc17a6 mRNA level cannot be detected by traditional in situ hybridization.

Characterization of real time co-expression of Vglut1, Vglut2 and Vglut3 in brainstem NA neurons in adult mice by fluorescent in situ hybridization.

(A) Representative images of Slc17a7 (Vglut1), Slc17a6 (Vglut2), and Slc17a8 (Vglut3) and Dbh (DBH) double ISH in coronal sections of WT mice in all of brainstem NA nuclei in adult mice. DBH (green), Vglut1/2/3 (red), DAPI (cyan). Scale bar 50μm. (B) Quantification of the percentage of Slc17a6 (Vglut2) co-expression in NA neurons in each NA nuclei. Female data (pink), male data (blue). (C) Quantification of the percentage of Slc17a8 (Vglut3) co-expression in NA neurons in each NA nuclei. Female data (pink), male data (blue).

Vglut2 expression is effectively knocked down in the whole central noradrenergic system in Dbh-Cre; Vglut2 cKO mice

To investigate the requirement of Vglut2-based glutamatergic signaling in the central NA system in breathing under physiological challenges, we used Dbh-Cre; Slc17a6flox/flox(Dbh-Cre; Vglut2 cKO) to remove Vglut2 expression from all noradrenergic neurons. This is the same model used by Abbott et al. (2014) to show that Vglut2 expression in the NA system was required for an increase in respiratory frequency following unilateral optogenetic stimulation of anterior C1. To verify if Vglut2 expression was effectively removed from central NA neurons in Dbh-Cre; Vglut2 cKO mice, we used fluorescent mRNA in situ hybridization for Slc17a6 (Vglut2) and Dbh in the mouse brainstems and compared the Slc17a6 (Vglut2) signal intensity in Dbh positive neurons in Dbh-Cre; Vglut2 cKO mice to their littermate controls. We found that Slc17a6 mRNA expression was 96.4 ± 0.226% decreased compared to controls, indicating that the Slc17a6 (Vglut2) was effectively recombined to abrogate expression in NA neurons (Figure 3A-B and Figure 3 supplement 1).

DBH with Vglut2 co-staining by fluorescent in situ hybridization confirms that Vglut2 expression is disrupted in the whole central NA system in Dbh-Cre; Vglut2 cKO mice.

(A) Representative images of Slc17a6 (Vglut2) and Dbh (DBH) double ISH in anterior C1/A1, posterior C1/A1, anterior C2/A2 and posterior C2/A2 in control and Dbh-Cre; Vglut2 cKO mice. DBH (green), Vglut2 (red), DAPI (cyan). Scale bar 20μm.(B) Quantification of intensity of Slc17a6 (Vglut2) signal in control and Dbh-Cre; Vglut2 cKO mice.

Vglut2-based glutamatergic signaling in central NA neurons is not required for baseline breathing nor for the hypercapnic ventilatory reflex under distinct CO2 challenges

Earlier version: Multiple prior studies provide indirect evidence that NA glutamate transmission may play roles in respiratory function, particularly that anterior C1 neurons mediate the hypoxic ventilatory response through the pFRG/RTN and/or the NA A5 group (Malheiros-Lima et al., 2022, 2020). Additionally, the A5 group has been implicated in the hypercapnic reflex (Haxhiu et al., 1996). Thus, we sought to determine if Vglut2-based glutamatergic signaling in central NA neurons is necessary to regulate baseline breathing and the hypercapnic chemoreflex in unanesthetized and unrestrained animals. We used whole-body barometric plethysmography to measure the breathing and metabolism of unanesthetized and unrestrained mice. After a five-day habituation protocol, the respiratory responses of Dbh-Cre; Vglut2 cKO mice and their littermate controls, including the overall respiratory output or ventilatory equivalent of oxygen (VE/VO2), respiratory rate (Vf), tidal volume (VT), minute ventilation (VE), metabolic demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE), were measured under room air and hypercapnia (5% CO2) (Figure 4A). 5% CO2 gas challenge produced a significant change in all the respiratory parameters compared to room air in both Dbh-Cre; Vglut2 cKO and control mice. However, surprisingly, Dbh-Cre; Vglut2 cKO mice did not show significant differences in any respiratory or metabolic parameter mentioned above compared to their littermate controls under both room air and 5% CO2 challenges (Figure 4B and Table 1-2). In addition, to investigate if Vglut2-based NA derived glutamate is required for regulating the dynamic patterns of breathing, such as the breathing rhythms, we generated Poincaré plots for TI, TE, TTOT, VT/TI, VT/TE, Vf, VT, and VE and we calculated the SD1 and SD2 statistics for each parameter under room air and 5% CO2. We found that Dbh-Cre; Vglut2 cKO mice only showed a marginally statistically significant increase in the SD1 of inspiratory duration (TI) (p=0.0492). Outside of this single comparison, however, we found no significant differences in the other parameters between the two groups (Figure 4 supplement 4A-H). Thus, we theorize that the statistically significant result for the SD1 of TI seen here is likely either a result of measurement noise or a consequence of a type 1 error. The overall pattern of results suggested that Vglut2-based glutamatergic signaling is not required for regulating either the steady state of or the dynamic patterns of baseline breathing and hypercapnic ventilatory reflex under 5% CO2.

Revised version: Multiple prior studies provide indirect evidence that NA glutamate transmission may play roles in respiratory function, particularly that anterior C1 neurons mediate the hypoxic ventilatory response through the pFRG/RTN and/or the NA A5 group (Malheiros-Lima et al., 2022, 2020). Additionally, the A5 group has been implicated in the hypercapnic reflex (Haxhiu et al., 1996). Thus, we sought to determine if Vglut2-based glutamatergic signaling in central NA neurons is necessary to regulate baseline breathing and the hypercapnic chemoreflex in unanesthetized and unrestrained animals. We used whole-body barometric plethysmography to measure the breathing and metabolism of unanesthetized and unrestrained mice. After a five-day habituation protocol, the respiratory responses of Dbh-Cre; Vglut2 cKO mice and their littermate controls, including the overall respiratory output or ventilatory equivalent of oxygen (VE/VO2), respiratory rate (Vf), tidal volume (VT), minute ventilation (VE), metabolic demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE), were measured under room air and hypercapnia (5% CO2) (Figure 4A). 5% CO2 gas challenge produced a significant change in all the respiratory parameters compared to room air in both Dbh-Cre; Vglut2 cKO and control mice. However, surprisingly, Dbh-Cre; Vglut2 cKO mice did not show significant differences in any respiratory or metabolic parameter mentioned above compared to their littermate controls under both room air and 5% CO2 challenges (Figure 4B and Table 1-2). In addition, to investigate if Vglut2-based NA derived glutamate is required for regulating the dynamic patterns of breathing, such as the breathing rhythms, we generated Poincaré plots for TI, TE, TTOT, VT/TI, VT/TE, Vf, VT, and VE and we calculated the SD1 and SD2 statistics for each parameter under room air and 5% CO2. We found that Dbh-Cre; Vglut2 cKO mice only showed a marginally statistically significant increase in the SD1 of inspiratory duration (TI) (p=0.0492). Outside of this single comparison, however, we found no significant differences in the other parameters between the two groups (Figure 4 supplement 1A-H). Thus, we theorize that the statistically significant result for the SD1 of TI seen here is likely either a result of measurement noise or a consequence of a type 1 error. The overall pattern of results suggested that Vglut2-based glutamatergic signaling is not required for regulating either the steady state of or the dynamic patterns of baseline breathing and hypercapnic ventilatory reflex under 5% CO2.

Vglut2 conditional knockout in central noradrenergic neurons fails to alter baseline breathing and the hypercapnic ventilatory reflex (5% CO2).

(A) Mouse model schematic and experimental protocol including habituation and hypercapnia protocol (5% CO2). (B) Under both room air and hypercapnia (5% CO2), Dbh-Cre; Vglut2 cKO mice didn’t show significant changes in respiratory output (VE/VO2), minute ventilation (VE), respiratory rate (Vf), tidal volume (VT), metabolism demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE). Linear mixed-effects regression model, ****p<0.0001, **p<0.01, *p<0.05, ns: p20.05.

To further investigate the role of Vglut2-based glutamatergic signaling in the hypercapnic response, we interrogated the system further with more severe hypercapnic exposures and measured the breathing of Dbh-Cre; Vglut2 cKO mice under 7% CO2 and 10% CO2. Both 7% and 10% CO2 stimuli produced a significant respiratory response in both mutant and control mice. Under the 7% CO2 condition, Dbh-Cre; Vglut2 cKO mice showed a significantly decreased tidal volume (VT) (p=0.023) and a significantly reduced inspiratory flow rate (VT/TI) (p=0.0016), but the overall respiratory output (VE/VO2), respiratory rate (Vf), minute ventilation (VE), metabolic demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), and expiratory flow rate (VT/TE) did not show a significant difference between the mutant mice and their sibling controls (Figure 5A-B and Table 3). Additionally, none of the dynamic patterns of any respiratory parameter showed a significant difference between the mutant and control groups (Figure 5 supplement 1A-H). This result suggests that Vglut2-based glutamatergic signaling in central NA neurons is not required for regulating the hypercapnic ventilatory reflex and breathing regularity under 7% CO2.

Vglut2 conditional knockout in central noradrenergic neurons fails to alter the majority of hypercapnic ventilation parameters and metabolism under 7% CO2.

(A) Mouse model schematic and experimental protocol including habituation and hypercapnia protocol (7% CO2). (B) Knocking out Vglut2 in the whole NA system failed to alter VE/VO2, Vf, VE, VO2, TI, TE, TTOT, VT/TE but only showed a significantly reduce in VT and VT/TI under 7% CO2 condition. Linear mixed-effects regression model, ****p<0.0001, **p<0.01, *p<0.05, ns: p20.05.

Under 10% CO2 challenge, none of the steady state breathing parameters including the overall respiratory outputs (VE/VO2), respiratory rate (Vf), tidal volume (VT), minute ventilation (VE), metabolic demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE) were significantly different between Dbh-Cre; Vglut2 cKO mice and their littermate controls, suggesting that Vglut2-based glutamatergic signaling in central NA neurons is not required for the hypercapnic ventilatory reflex under 10% CO2 (Figure 6A-B and Table 4). For the dynamic patterns of breathing, the SD1 of breath cycle duration (TTOT) (p=0.028) and SD1 of respiratory rate (Vf) (p=0.038) showed a significant increase in mutant mice compared to the controls (Figure 6 supplement 1A-H). This suggests that removing Vglut2 from central NA neurons may increase the breathing irregularity under high CO2 challenges like 10% CO2.

Vglut2 conditional knockout in central noradrenergic neurons fails to alter the hypercapnic ventilatory reflex (10% CO2).

(A) Mouse model schematic and experimental protocol including habituation and hypercapnia protocol (10% CO2). (B) Knocking out Vglut2 in the whole NA system didn’t significantly alter VE/VO2, Vf, VT, VE, VO2, TI, TE, TTOT, VT/TI, and VT/TE under 10% CO2 condition. Linear mixed-effects regression model, ****p<0.0001, **p<0.01, *p<0.05, ns: p20.05.

Vglut2-based glutamatergic signaling in central NA neurons is not required for the hypoxic ventilatory reflex under 10% O2

To determine if Vglut2-based glutamatergic signaling in central NA neurons is necessary to regulate the hypoxic chemoreflex, we measured the ventilation response of Dbh-Cre; Vglut2 cKO mice under 10% O2. We analyzed the ten respiratory parameters in three 5-min epochs following 10% O2 challenge onset. Similar to the hypercapnic challenges, the 10% O2 challenge produced a significant breathing response in every 5-min epoch compared to room air, but no significant difference was found in the overall respiratory output (VE/VO2), respiratory rate (Vf), tidal volume (VT), minute ventilation (VE), metabolism demand (VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE) between Dbh-Cre; Vglut2 cKO mice and controls in any of the three 5-min 10% O2 epochs (Figure 7A-B and Table 5). These data suggest that Vglut2-based glutamatergic signaling in central NA neurons is not required for the hypoxic ventilatory reflex under 10% O2. For the dynamic patterns of breathing, the SD2 of inspiratory flow rate (VT/TI) (p=0.005) and the SD2 of minute ventilation (VE) (p=0.026) were significantly increased during the first epoch (first 5 minutes) under the 10% O2 challenge in mutant mice, suggesting that Vglut2-based glutamate may play a role in maintaining breathing regularity during the initial, or reflexive, response to hypoxia (Figure 7 supplement 1A-H).

Vglut2 conditional knockout in central noradrenergic neurons fails to alter the hypoxic ventilatory reflex (10% O2).

(A) Mouse model schematic and experimental protocol including habituation and hypoxia protocol (10% O2). (B) Knocking out Vglut2 in the whole NA system did not show significant breathing changes including VE/VO2, VE, Vf, VT, VO2, TI, TE, TTOT, VT/TI, and VT/TE in either of the three five-minute time periods under hypoxia (10% O2). Linear mixed-effects regression model, ****p<0.0001, **p<0.01, *p<0.05, ns: p20.05.

Discussion

Neuron co-transmission of two or more neurotransmitters across excitatory, inhibitory, and neuro-modulatory facets is becoming an increasingly appreciated phenomenon in the central nervous system. Functional interrogation of co-transmission is essential not only to understand the distinct and cooperative roles that multiple signaling molecules may play at the synapse but also to determine the most critical targets for potential therapeutics. To our knowledge, the unopposed paradigm in the field until now has been that vesicular glutamate transporter 2 (Vglut2)-based glutamate transmission from central noradrenergic neurons is important in respiratory homeostasis. Anatomically, it has long been appreciated that central NA neurons co-express Vglut2, the major glutamate marker among three glutamate transporters (DePuy et al., 2013; Souza et al., 2022a; Stornetta et al., 2002a, 2002b; Yang et al., 2021). Also, it has been well documented that NA Vglut2 positive fibers innervate known central respiratory centers or other autonomic centers, that, when perturbed, could potentially disrupt respiratory homeostasis (Supplemental Table 1). Functionally, multiple reports have suggested that Vglut2-based glutamate transmission plays a role in respiratory homeostasis (Abbott et al., 2014; Malheiros-Lima et al., 2022, 2020, 2018). It has also been suggested that, at least for C1 neurons, the apparent lack of a plasmalemmal monoamine transporter may attenuate or eliminate noradrenergic or adrenergic release from these fibers (Guyenet et al., 2013). However, as discussed below, much of the functional evidence supporting a role for Vglut2-based glutamate transmission from central noradrenergic neurons in breathing is circumstantial or of a non-physiological nature, and absent or attenuated release of adrenaline and noradrenaline has not been demonstrated. Nonetheless, the dominant perspective that NA-Vglut2 glutamate transmission plays a role in breathing has remained unchallenged. Our studies, in contrast to prior work, show that loss of Vglut2 in NA neurons does not appreciably change baseline or chemosensory breathing in unanesthetized and unrestrained mice. Additionally, our work uncovers a novel dynamic expression pattern for Vglut2 and an entirely undescribed co-expression domain for Vglut3 in central NA neurons.

Vglut2 has been shown to be co-expressed in subsets of central NA neurons, however, to our knowledge, the potential for the central NA neurons to express the other glutamate transporters Vglut1 or Vglut3 has not been reported at the anatomical level. To determine if central NA neurons express either Vglut1 or Vglut3 and to confirm the Vglut2 expression, we carried out cumulative intersectional fate mapping using Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG and Slc17a8Cre; Dbhp2a-Flpo; Rosa26RC::FLTG compound mouse lines. As this approach maps the cumulative history of expression, we also examined acute adult expression of Slc17a7 (Vglut1), Slc17a6 (Vglut2) and Slc17a8 (Vglut3) using fluorescent RNA in situ hybridization. Our in situ results for Slc17a6 (Vglut2) agreed with earlier published outcomes (DePuy et al., 2013; Stornetta et al., 2002a, 2002b) showing as much as 80% expression in posterior NA neurons. However, there was no appreciable signal in anterior groups (A7, LC, A5, sub CD/CV). This lack of expression was in stark contrast to the cumulative fate map for Vglut2, showing as much as 75% of the anterior central NA neurons were recombined by the Vglut2-Cre over the lifetime of the animal. To further interrogate the acute activity of the Vglut2-Cre in the Locus Coeruleus in the adult mouse, we injected a Cre responsive virus in LC of Slc17a6Cre mice and found sparse Vglut2-Cre positive NA neurons in the LC. The outcomes from the cumulative fate map and acute injections are in agreement with published results from Yang et al. (2021). Notably, our use of the Dbhp2a-Flpoavoided expression in the parabrachial nucleus (PBN), which was a confounding variable in the ThFlpo mouse. In addition, 28.5% of A5 neurons were shown to be Slc17a6 (Vglut2) positive in adult rats by RNA scope (Souza et al., 2022b). It is notable that both our lab and the work published by (Stornetta et al., 2002a, 2002b) failed to detect Slc17a6 (Vglut2) expression in anterior groups, but that apparently very low levels can be detected by more sensitive recombinase and RNA scope methods. Additional electrophysiological evidence suggests that this expression is functional (Yang et al., 2021). However, it is notable that we see nearly 80% of the LC and 85% of A5 recombined in the intersectional strategy representing lifetime expression, but our study and others find only 20% of the LC and 28.5% of A5 may be expressing Slc17a6 (Vglut2) at any given moment in adults. Thus, it suggests that Vglut2 co-expression in anterior NA groups is time dependent, though the underlying temporal dynamics of Vglut2 remain unknown. We hypothesize that this could be either a temporal restriction in expression as the animal matures or may reflect dynamic change driven by behavioral or physiological experiences, which could be tested with the use of inducible conditional strategies.

For Slc17a7 (Vglut1), neither fate map nor in situ hybridization revealed any co-expression within the central NA system. However, Slc17a8 (Vglut3) expression was found in the posterior part of NA neurons including C2/A2 and C1/A1. This study is the first to characterize this expression. The majority of this Vglut3 co-expression is restricted to the posterior part of the C2/A2 NA population. This NA region has not been heavily implicated in respiratory control nonetheless it readily lends itself to interrogation by intersectional genetic methods and in situ hybridization. Given its distal expression from anterior C1, it is less likely that Vglut3 would compensate for loss of Vglut2 in anterior C1, though this has not been ruled out. Further, while Vglut3 is a marker of glutamatergic neurons, its role in neurotransmission is not wholly defined as Vglut3 is typically found in soma and dendrites (Fremeau et al., 2004).

The cumulative fate map and in situ hybridization experiments were performed in both females and males and no obvious sex difference was observed in the expression pattern of Vglut1, Vglut2 or Vglut3. However, only three mice were characterized in each dataset, thus, larger sample sizes may capture more subtle outcomes for potential sex differences.

Based on the Vglut2 co-expression in NA neurons and their projections to multiple brain regions important in respiration, several studies suggest that glutamate in NA populations is important in breathing control. However, the evidence those studies provided is circumstantial, indirect, or limited due to experimental caveats or technical limitations. Abbott et al. (2014), using exactly the same mouse model as this study, showed that conditionally knocking out Vglut2 in the whole central NA system did not significantly change respiratory rate under room air, which is consistent with our findings here. Our results differ in that they reported removing Vglut2 from central NA neurons attenuated an increase in respiratory rate resulting from unilateral C1 optogenetic stimulation. However, while the optogenetic stimulation of C1 neurons did drive a Vglut2 dependent increase in respiratory rate, it is not clear if the resulting change in breathing reflects a native feature of the network that might be engaged in chemosensory responses or other aspects of breathing. It remains possible that strong unilateral optogenetic stimulation (up to 20Hz) of a small, isolated population drives a Vglut2 dependent, but ectopic, function that is not typically engaged in the normal operation of the breathing network. Such ectopic outcomes could come from several potential, non-exclusive mechanisms. First, unilateral stimulation could drive a network hysteresis or dysregulation similar to a focal injury that is not a typical biological feature (Ducros et al., 2003). This point is well made in Abbott et al. (2013) in their discussion of C1 optogenetic stimulation where they state, “resulting cardiorespiratory response pattern should not be considered strictly “physiological,” because a physiological response is never initiated by selective activation of a single cluster of CNS neurons and the various subsets of C1 neurons are presumably never recruited en bloc under any physiological condition”. Second, strong stimulation could result in overwhelming synaptic mechanisms (i.e., glial uptake) resulting in glutamate spillover and unintended cross talk to affect extra-synaptic neurons proximal to targeted NA fibers. Activity dependent spillover transmission has been documented elsewhere in the nervous system (Henneberger et al., 2020; Hülsmann et al., 2000). Third, the outcomes could be secondary to a different autonomic or behavioral function requiring NA-derived glutamatergic signaling involving metabolism or circulation (DePuy et al., 2013; Yang et al., 2021). Additionally, respiratory parameters except for respiratory rate (Vf), such as tidal volume (VT) and minute ventilation (VE) were not reported, making comprehensive comparison difficult.

More circumstantially, Malheiros-Lima et al. (2020) suggested C1 neurons potentially release glutamate at the pFRG site to regulate active expiration under hypoxia in anesthetized rats by providing three indirect lines of evidence: 1) Vglut2-expressing C1 neurons project to the pFRG region; 2) Increased abdominal expiratory nerve activity (AbdEMG) was blunted after blockade of ionotropic glutamatergic receptors at the pFRG site under anesthesia using cytotoxic KCN mediated hypoxia; and 3) Depletion of C1 neurons eliminated the increased AbdEMG elicited by hypoxia. However, 1) the glutamatergic signaling targeting the pFRG region is not necessarily from C1 NA neurons since glutamatergic neurons from other regions also project to the pFRG site (Yang and Feldman, 2018). Thus, it is possible that the glutamatergic signaling that is required for hypoxic response is derived from other glutamatergic neurons such as the preBötzinger complex and the RTN itself. Furthermore, 2) C1 neuron depletion knocks out all the signaling modalities in the C1 population including both noradrenaline and glutamate. It is not clear which signaling drives the hypoxic response. Finally, 3) KCN injection mimics a hypoxic challenge by activating peripheral chemoreceptors, but the physiological responses differ from those seen during environmental hypoxia exposure. Thus, the abdominal activity changes due to KCN injection may not be recaptured under physiological hypoxic challenge.

Similarly, Malheiros-Lima et al. (2022) and Malheiros-Lima et al. (2018) showed that Vglut2-expressing C1 neurons project to both the NA A5 region and the preBötzinger complex. The blockade of ionotropic glutamatergic receptors at the A5 region or preBötzinger complex reduced the increase in phrenic nerve activity and the breathing frequency caused by optogenetic stimulation of C1 cells in an anesthetized preparation. Together these results suggest that Vglut2-expressing C1 neurons communicate with A5 or preBötzinger complex neurons by releasing glutamate in turn to regulate phrenic nerve activity or breathing frequency. Again, these two studies showed that C1 neurons are important to regulate phrenic nerve activity or breathing frequency, however, it is not necessarily through a direct C1-A5 or C1-preBötzinger complex glutamatergic pathway, whereas an indirect multi-synaptic pathway originating with C1 noradrenaline and ending with glutamate transmission from an intermediate to A5 or preBötzinger complex cannot be ruled out. Importantly, in all four studies noted here, respiratory measurements were limited or proxies of breathing under anesthesia were used. The role of NA-Vglut2 signaling was not reported in the context of an unanesthetized, unrestrained animal under normoxic, hypoxic, or hypercapnic breathing in these previous studies.

Thus, the limitations and indirect lines of the evidence in these studies raised the question as to whether or not glutamatergic signaling in central NA neurons is required or necessary to modulate respiratory homeostasis under more physiologically relevant circumstances and if NA-derived glutamate is an important mechanism to understand the pathophysiology of NA neurons in respiratory diseases. The goal of our study was to test the requirement of NA-based glutamatergic signaling in breathing homeostasis in the unanesthetized and unrestrained animal under conditions that would engage the whole respiratory control network rather than a singular component that may drive hysteresis or artifactual outcomes. Our in vivo breathing data under room air, hypoxic, and multiple hypercapnic conditions failed to show the requirement of NA-derived Vglut2 in normal breathing and respiratory chemoreflexes directly. For the dynamic patterns of breathing, we measured the SD1 and SD2 of eight parameters under each gas condition. We found no consistent change across multiple hypercapnic challenges or hypoxia. We did find that under severe hypercapnic challenge (10% CO2), NA derived Vglut2 removal resulted in an increase in the SD1 of breath cycle duration and respiratory rate, which was not observed under 5% and 7% CO2 conditions. Under hypoxia (10% O2), Dbh-Cre; Vglut2 cKO mice only showed an increase in the SD2 of inspiratory flow rate and minute ventilation during the first 5-minute of 10% O2 exposure. These findings may suggest that Vglut2-based glutamatergic signaling in central NA neurons may play a role in modulating breathing regularity under high CO2 challenge and hypoxia. However, the phenotypes were of a small magnitude and not consistent across challenges, and not at all seen under homeostatic breathing.

To ensure a rigorous and robust conclusion, we included several extra precautions and additional measures in our experimental design and analysis. First, we use exactly the same mouse model, DBH-Cre; Vglut2 cKO, as Abbott et al. (2014) used, and confirmed, by a second method (fluorescent RNA in situ hybridization) that Vglut2 expression was indeed abrogated in the whole central NA system. Our results are in strong agreement with DePuy et al. (2013) that recombination is efficient. Notably, we assayed for expression at the cell body, rather than the fibers and see no singular cell body that is scored positive. Second, we measured the breathing of unanesthetized, unrestrained, and habituated mice under physiological challenges including normoxia, hypercapnia and hypoxia. Third, to visualize breathing function comprehensively, the metabolic parameter and multiple respiratory parameters (both the steady state and the dynamic patterns) were measured and reported in absolute values including oxygen consumption, respiratory rate, tidal volume, minute ventilation, overall respiratory output (minute ventilation normalized to oxygen consumption), inspiratory duration, expiratory duration, breath cycle duration, inspiratory flow rate, and expiratory flow rate. Fourth, our experimental design was overpowered (i.e., n=16-21 vs 5-13 needed for power based on power analysis and n=7-8 reported in other manuscripts). Each respiratory/metabolic parameter for each physiological condition including room air, hypercapnia, and hypoxia was derived by quantifying the entire respiratory trace and analyzed between mutant and sibling controls by using a linear mixed-effects regression model with animal type (experimental vs. control) as fixed effects, animal ID as a random effect, and sex as an experimental covariate (for which we did not find a significant effect) (Lusk et al., 2023).

Earlier version: Despite our experimental design and the extensive nature of our measurements, there remain several possibilities or considerations in the interpretation of our data. The first possibility is Vglut3 in posterior CA domain (first characterized in this study) compensates for the effect of the loss of Vglut2 in NA neurons in breathing. However, we argue that this is unlikely because 1) Our model is the same used in the other study where a Vglut2-dependent effect was seen with optogenetic stimulation. 2) Vglut3 is expressed in the posterior C2/A2 NA neurons which are anatomically distant from anterior C1 and the LC (the two most likely NA candidates with Vglut2 co-expression to impact breathing). Also, posterior C2/A2 hasn’t been heavily implicated in breathing regulation. Future work will need to further investigate the role of Vglut3 in posterior CA in respiratory control in order to test this possibility. The second possibility is that pre- or post-natal developmental compensation corrects for the loss of NA derived glutamate. Again, we argue that is unlikely, as 1) we are using the same adult model that provided the most direct evidence and first functionally demonstrated a potential role for C1-derived glutamate in breathing; 2) Abbott et al. (2014) highlight that C1 NA neurons showed no obvious abnormalities in number, morphology, and projection patterns after knocking out Vglut2; and 3) we have examined the requirement of NA-expressed Vglut2 in P7 neonate mice in the autoresuscitation reflex and saw no differences (data not shown). However, it is still possible that developmental compensation occurs before P7, which is something that could be better tested by using the Dbh-CreERT2 (Stubbusch et al., 2011) to more acutely remove Vglut2 expression (though this leaves a window of 2-3 weeks for compensation to occur). Third, as we only address the loss of NA-Vglut2 based signaling across the entire NA system, it remains a formal possibility that this ultimately results in the removal of Vglut2 dependent signaling from two counter balancing regions of the NA system and are therefore left with null results whereas removal from anterior C1 alone would show a requirement. To test this possibility, it would require an intersectional conditional knockout/loss of function approach, something that has not yet been shown to be efficient and effective and would be beyond the scope of these studies. Fourth, it is possible that another neurotransmitter/neuropeptide within NA neurons might be released in higher amounts in Dbh-Cre; Vglut2 cKO mice to compensate for the deficiency of glutamate in breathing. Neuropeptide Y (NPY) and galanin have been shown to co-exist in LC NA neurons and are involved in regulating energy metabolism, stress, or anxiety (Ruohonen et al., 2012; Tasan et al., 2010; Weinshenker and Holmes, 2016). In our studies, we account for potential metabolic changes that may underlie or be coordinate with changes in breathing. However, we did not test for potential changes in blood pressure or other autonomic functions as well as affective state that may, in some obscure way, compensate for changes in breathing. Further interrogation of NPY or galanin signaling across NA neurons and potential compensatory effects in breathing through changes in metabolism or other autonomic function under stress or anxiety may yield notable insights. Lastly, our results are loss of function in nature. We do not test sufficiency or gain of function and cannot fully rule out a role for NA-Vglut2 based glutamatergic signaling in the control of breathing.

Revised version: Despite our experimental design and the extensive nature of our measurements, there remain several possibilities or considerations in the interpretation of our data. The first possibility is Vglut3 in posterior CA domain (first characterized in this study) compensates for the effect of the loss of Vglut2 in NA neurons in breathing. However, we argue that this is unlikely because 1) Our model is the same used in the other study where a Vglut2-dependent effect was seen with optogenetic stimulation. 2) Vglut3 is expressed in the posterior C2/A2 NA neurons which are anatomically distant from anterior C1 and the LC (the two most likely NA candidates with Vglut2 co-expression to impact breathing). Also, posterior C2/A2 hasn’t been heavily implicated in breathing regulation. Future work will need to further investigate the role of Vglut3 in posterior CA in respiratory control in order to test this possibility. The second possibility is that pre- or post-natal developmental compensation corrects for the loss of NA derived glutamate. Again, we argue that is unlikely, as 1) we are using the same adult model that provided the most direct evidence and first functionally demonstrated a potential role for C1-derived glutamate in breathing; 2) Abbott et al. (2014) highlight that C1 NA neurons showed no obvious abnormalities in number, morphology, and projection patterns after knocking out Vglut2; and 3) we have examined the requirement of NA-expressed Vglut2 in P7 neonate mice in the autoresuscitation reflex and saw no differences (data not shown). However, it is still possible that developmental compensation occurs before P7, which is something that could be better tested by using the Dbh-CreERT2 (Stubbusch et al., 2011) to more acutely remove Vglut2 expression (though this leaves a window of 2-3 weeks for compensation to occur). Third, as we only address the loss of NA-Vglut2 based signaling across the entire NA system, it remains a formal possibility that this ultimately results in the removal of Vglut2 dependent signaling from two counter balancing regions of the NA system and are therefore left with null results whereas removal from anterior C1 alone would show a requirement. To test this possibility, it would require an intersectional conditional knockout/loss of function approach, something that has not yet been shown to be efficient and effective and would be beyond the scope of these studies. Fourth, it is possible that another neurotransmitter/neuropeptide within NA neurons might be released in higher amounts in Dbh-Cre; Vglut2 cKO mice to compensate for the deficiency of glutamate in breathing. Loss of Vglut2 could reduce dopamine release in subsets of dopaminergic neurons (Alsiö et al., 2011; Fortin et al., 2012; Hnasko et al., 2010). Thus, it remains possible that loss of Vglut2 affects dopamine loading into the vesicles and in turn affecting noradrenaline release as Dopamine Beta Hydroxylase (DBH) synthesizes noradrenaline from dopamine inside the vesicles (Weinshenker, 2007). Changes in noradrenaline release in Vglut2 negative NA neurons could be further examined with fast-scan cyclic voltammetry (FSCV) or microdialysis. Neuropeptide Y (NPY) and galanin have been shown to co-exist in LC NA neurons and are involved in regulating energy metabolism, stress, or anxiety (Ruohonen et al., 2012; Tasan et al., 2010; Weinshenker and Holmes, 2016). In our studies, we account for potential metabolic changes that may underlie or be coordinate with changes in breathing. However, we did not test for potential changes in blood pressure or other autonomic functions as well as affective state that may, in some obscure way, compensate for changes in breathing. Further interrogation of NPY or galanin signaling across NA neurons and potential compensatory effects in breathing through changes in metabolism or other autonomic function under stress or anxiety may yield notable insights. Lastly, our results are loss of function in nature. We do not test sufficiency or gain of function and cannot fully rule out a role for NA-Vglut2 based glutamatergic signaling in the control of breathing.

From a translational perspective, our data question whether or not glutamatergic signaling in NA neurons is likely to be a key mechanism and therefore a therapeutic target for breathing disorders, such as Rett syndrome and SIDS. Mecp2-deficient mice (a mouse model of Rett Syndrome that phenocopies many human symptoms) showed both a deficiency of NA populations (reduced number of NA neurons in C2/A2 and C1/A1 group) and highly variable respiratory rhythm at around 4-5 weeks of age (Roux et al., 2007; Viemari et al., 2005). 80% of Rett Syndrome patients experience breathing issues, such as unstable breathing, episodes of hyperventilation, and breath holds, throughout their lifespan (Ramirez et al., 2020). SIDS decedents show NA abnormalities and are hypothesized to ultimately succumb from a failure in cardiorespiratory autoresuscitation (Chigr et al., 1989; Garcia et al., 2013; Kopp et al., 1993; Mansouri et al., 2001; Ozawa et al., 2003; Takashima and Becker, 1991). However, Dbh-Cre; Vglut2 cKO mice, have normal baseline and chemosensory respiratory parameters as adults and neonate mice show normal autoresuscitation indicating that perturbed NA based glutamatergic signaling may not be a key driver in these or other related respiratory pathophysiologies.

In conclusion, our studies show that Vglut2 has a dynamic and extensive expression profile across the central NA system. We were able to characterize a novel Vglut3-expressing NA population in the posterior C2/A2 nuclei. Despite prior studies providing indirect and circumstantial evidence that noradrenergic-based glutamatergic signaling may play a role in control of breathing, our conditional loss of function studies in adult unanesthetized and unrestrained mice failed to consistently and significantly change room air breathing, the hypercapnic ventilatory reflex, and the hypoxic ventilatory reflex. These outcomes offer a contrasting perspective from the current field view and provides further insight into the potential role of NA-glutamate transmission in the control of breathing and NA neuron dysfunction in respiratory disorders.

Materials and methods

Ethical approval

Studies were approved by the Baylor College of Medicine Institutional Animal Care and Use Committee (IACUC) under protocol AN-6171, and all experiments reported here were performed in accordance with relevant guidelines and regulations.

Breeding, genetic background, and maintenance of mice

Earlier version: We maintained all our heterozygous mouse strains by backcrossing to wildtype C57BL/6J mice and homozygous mouse strains by sibling crosses. For immunofluorescence experiments, heterozygous B6;129S-Slc17a7tm1.1(cre)Hze/J (Slc17a7Cre, Vglut1-Cre) (Jax Stock No: 023527), Slc17a6tm2(cre)Lowl/J (Slc17a6Cre, Vglut2-Cre) (Jax Stock No: 016963) and B6;129S-Slc17a8tm1.1(cre)Hze/J (Slc17a8Cre, Vglut3-Cre) (Jax Stock No: 018147) were mated with heterozygous Dbhem2.1(flpo)Rray (Dbhp2a-Flpo, DBH-p2a-Flpo) (MMRRC ID: 41575) (Sun and Ray, 2016) mice respectively to derive compound lines with both Cre and Flpo alleles (Slc17a7Cre; Dbhp2a-Flpo, Slc17a6Cre; Dbhp2a-Flpo, and Slc17a8Cre; Dbhp2a-Flpo). Then these three compound lines were each mated with homozygous B6.Cg-Gt(ROSA)26Sortm1.3(CAG-tdTomato,-EGFP)Pjen/J (Rosa26RC::FLTG) (Jax Stock No: 026932) to derive three different intersectional mouse lines Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTG and Slc17a8Cre; Dbhp2a-Flpo; Rosa26RC::FLTG. For in situ hybridization experiments, wildtype C57BL/6J mice were ordered from the Center of Comparative Medicine (CCM), Baylor College of Medicine. For NA-Vglut2 conditional loss of function in situ hybridization and plethysmography experiments, hemizygous transgene Tg(Dbh-cre)KH212Gsat (Dbh-Cre) (MMRRC ID: 036778-UCD GENSAT) mice were mated with homozygous Slc17a6tm1Lowl/J (Slc17a6flox/flox) (Jax Stock No: 012898) to derive Dbh-Cre; Slc17a6flox/+. Dbh-Cre; Slc17a6flox/+mice were mated with Slc17a6flox/flox to derive Dbh-Cre; Slc17a6flox/flox(Dbh-Cre; Vglut2 cKO). Sibling mice that lacked the Cre allele or carried the Cre allele but lacked the floxed Vglut2 alleles were used as controls. Rosa26 specific primers for the Rosa26RC::FLTG mice were 5’-GCACTTGCTCTCCCAAAGTC, 5’-GGGCGTACTTGGCATATGAT, and 5’-CTTTAAGCCTGCCCAGAAGA (Ray et al., 2011) and yield a 495 bp band (targeted) and 330 bp band (wt). Cre-specific primers for all Cre drivers were 5’-ATCGCCATCTTCCAGCAGGCGCACCATTGCCC and 5’-GCATTTCTGGGGATTGCTTA and yielded a 550 bp band if positive. Flpo-specific primers for Dbhp2a- Flpo are 5’-CACGCCCAGGTACTTGTTCT and 5’-CCACAGCAAGAAGATGCTGA (Sun and Ray, 2016) and yielded a 226 bp band if positive. Slc17a6flox specific primers for Vglut2-floxed mice are available at The Jackson Laboratory website (https://www.jax.org/strain/012898).

Revised version: We maintained all our heterozygous mouse strains by backcrossing to wildtype C57BL/6J mice and homozygous mouse strains by sibling crosses. For immunofluorescence experiments, heterozygous B6;129S-Slc17a7tm1.1(cre)Hze/J (Slc17a7Cre, Vglut1-Cre) (Jax Stock No: 023527), Slc17a6tm2(cre)Lowl/J (Slc17a6Cre, Vglut2-Cre) (Jax Stock No: 016963) and B6;129S-Slc17a8tm1.1(cre)Hze/J (Slc17a8Cre, Vglut3-Cre) (Jax Stock No: 018147) were mated with heterozygous Dbhem2.1(flpo)Rray (Dbhp2a-Flpo, DBH-p2a-Flpo) (MMRRC ID: 41575) (Sun and Ray, 2016) mice respectively to derive compound lines with both Cre and Flpo alleles (Slc17a7Cre; Dbhp2a-Flpo, Slc17a6Cre; Dbhp2a-Flpo, and Slc17a8Cre; Dbhp2a-Flpo). Then these three compound lines were each mated with homozygous B6.Cg-Gt(ROSA)26Sortm1.3(CAG-tdTomato,-EGFP)Pjen/J (Rosa26RC::FLTG) (Jax Stock No: 026932) to derive three different intersectional mouse lines Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTGand Slc17a8Cre; Dbhp2a-Flpo; Rosa26RC::FLTG. For in situ hybridization experiments, wildtype C57BL/6J mice were ordered from the Center of Comparative Medicine (CCM), Baylor College of Medicine. For NA-Vglut2 conditional loss of function in situ hybridization and plethysmography experiments, hemizygous transgene Tg(Dbh-cre)KH212Gsat (Dbh-Cre) (MMRRC ID: 036778-UCD GENSAT) mice were mated with homozygous Slc17a6tm1Lowl/J (Slc17a6flox/flox) (Jax Stock No: 012898) to derive Dbh-Cre; Slc17a6flox/+. Dbh-Cre; Slc17a6flox/+mice were mated with Slc17a6flox/flox to derive Dbh-Cre; Slc17a6flox/flox(Dbh-Cre; Vglut2 cKO). Sibling mice that lacked the Cre allele or carried the Cre allele but lacked the floxed Vglut2 alleles were used as controls. Rosa26 specific primers for the Rosa26RC::FLTG mice were 5’-GCACTTGCTCTCCCAAAGTC, 5’-GGGCGTACTTGGCATATGAT, and 5’-CTTTAAGCCTGCCCAGAAGA (Ray et al., 2011) and yield a 495 bp band (targeted) and 330 bp band (wt). Cre-specific primers for all Cre drivers were 5’-ATCGCCATCTTCCAGCAGGCGCACCATTGCCC and 5’-GCATTTCTGGGGATTGCTTA and yielded a 550 bp band if positive. Flpo-specific primers for Dbhp2a- Flpo are 5’-CACGCCCAGGTACTTGTTCT and 5’-CCACAGCAAGAAGATGCTGA (Sun and Ray, 2016) and yielded a 226 bp band if positive. Slc17a6flox specific primers for Vglut2-floxed mice are available at The Jackson Laboratory website (https://www.jax.org/strain/012898).

Immunofluorescence staining

Slc17a7Cre; Dbhp2a-Flpo; Rosa26RC::FLTG, Slc17a6Cre; Dbhp2a-Flpo; Rosa26RC::FLTGand Slc17a8Cre; Dbhp2a- Flpo; Rosa26RC::FLTG adult mice of both sexes were sacrificed and transcardially perfused with 0.1 M phosphate-buffered saline (PBS) then with 4% paraformaldehyde (PFA) in PBS. Mouse brains were dissected out and fixed for 2h in 4% PFA before a PBS rinse and dehydration in 30% sucrose in PBS. Brains were embedded in OCT blocks, sectioned at the thickness of 30μm, mounted on slides, and stored at −80°C until they were ready for staining. The slides were hydrated in 0.1% Triton-X in PBS (PBST) for 15 mins, blocked with 5% donkey serum in 0.1% PBST for 1h at room temperature and then incubated with primary antibodies for 72 hours at 4°C in 0.1% PBST with 5% donkey serum. Tissues were washed in 0.1% PBST 3 times for 10 min each and then incubated with secondary antibodies for 2hs at room temperature in 0.1% PBST with 5% donkey serum. Slides were washed with 0.1% PBST for 10 min and washed in PBS twice for 10 min each, stained for DAPI, washed 3 times for 10 min each with PBS and mounted in ProLong Glass (Invitrogen). The following primary and secondary antibodies were used: chicken anti-GFP (1:1,000, Abcam ab13970), rabbit anti-dsRed (1:1,000, Clontech 632496), donkey anti-chicken Cy2 (1:500, Jackson 703-225-155), donkey anti-rabbit Cy3 (1:500, Jackson 711-165-152).

In situ hybridization

Mouse brains were dissected out from adult mice of both sexes with 6-8 week of age, sectioned into 25µm brain sections and mounted on slides. We generated a digoxigenin (DIG)-labeled mRNA antisense probe against Slc17a7 (Vglut1), Slc17a6 (Vglut2), and Slc17a8 (Vglut3) and fluorescein (FITC)-labeled mRNA against Dbh using reverse-transcribed mouse cDNA as a template and an RNA DIG or FITC-labeling kits from Roche (Sigma). Primer and probe sequences for the Slc17a7 (Vglut1), Slc17a6 (Vglut2), and Slc17a8 (Vglut3) and Dbh probes are available in the Allen Brain Atlas (http://www.brain-map.org). For the Slc17a6 (Vglut2) and Dbh double ISH in Dbh-Cre; Vglut2 cKO and their littermate controls, we generated a new Slc17a6 (Vglut2) probe targeting exon 2 of Slc17a6 specifically (Tong et al., 2007) and the probe sequence is 892-1144bp as Slc17a6 transcript variant 1 and the size is 253bp. ISH was performed by the RNA In Situ Hybridization Core at Baylor College of Medicine using an automated robotic platform as previously described (Yaylaoglu et al., 2005) with modifications of the protocol for double ISH. Modifications in brief [see Yaylaoglu et al. (2005) for buffer descriptions]: both probes were hybridized to the tissue simultaneously. After the described washes and blocking steps the DIG-labeled probes were visualized using tyramide-Cy3 Plus (1/75 dilution, 15-minute incubation, Akoya Biosciences). After washes in TNT, the remaining HRP-activity was quenched by a 10-minute incubation in 0.2M HCl. The sections were then washed in TNT, blocked in TNB for 15 minutes, and incubated at room temperature for 30 minutes with HRP-labeled sheep anti-FITC antibody (1/500 in TNB, Roche/Sigma). Following washes in TNT, the FITC-labeled probe was visualized using tyramide-FITC Plus (1/50 dilution, 15-minute incubation, Akoya Biosciences). Following washes in TNT, the slides were stained with DAPI (invitrogen), washed again, removed from the machine, and mounted in ProLong Diamond (Invitrogen).

Viral injection

To verify Vglut2 co-expression in LC in adult mice (Yang et al., 2021), 6-week-old Slc17a7Cre mice were injected with pAAV-EF1a-DIO-tdTomato-WPRE virus (RRID:Addgene_133786, obtained from Joshua Ortiz at the Optogenetics and Viral Vectors Core at the Jan and Dan Duncan Neurological Research Institute, 1ul at 7.90E+10 Gc/mL) into the LC (coordinates from bregma anteroposterior −5.4 mm, lateral + 0.8 mm and dorsoventral – 4.0 mm) and allowed to incubate for 4 weeks.

Plethysmography

Plethysmography on unanesthetized and free-moving mice was carried out as described in (Ray et al., 2011). 6–8-week-old adult mice of both sexes with minimum group sizes n=16 were used in both the experimental and control group for each experiment. Animals were subjected to a five-day habituation protocol with each day including several minutes of handling, temperature taken by rectal probe and at least 30 min exposure in the plethysmography chamber (Martinez et al., 2019). Plethysmography started to be performed on the sixth day and was finished for all animals within a week of the last day of habituation (Maximum 6 mice can be assayed for plethysmography per day due to the limited number of plethysmography rigs). On the day of testing, mice were taken out from their home cage, weighed, and rectal temperature was taken. Animals were then placed into a flow-through, temperature-controlled (about 32°C with real-time temperature recording) plethysmography chamber and allowed to acclimate for 20-40 mins in room air (21% O2/79% N2) conditions. After acclimation (indicated by a steady respiratory trace free from movement artifact), a 20-min baseline breathing trace was taken under room air, then the chamber gas was switched to a hypercapnic or hypoxic mixture of 5% CO2/21% O2/74% N2, 7% CO2/21% O2/72% N2, 10% CO2/21% O2/69% N2, or 10% O2/90% N2, depending on the protocol, for 20 min. Chamber gas was then switched back to room air for another 20 min. The animals were removed from the chamber and rectal temperature was measured immediately after the mice were taken out from the chamber. Each testing period for an individual mouse was separated by 24hs to allow for a full recovery.

Plethysmography data analysis and statistics

Details were previously described in (Martinez et al., 2019; Ray et al., 2011). Plethysmography pressure changes were measured using a Validyne DP45 differential pressure transducer, CD15 carrier demodulator and a reference chamber, and were recorded with LabChart Pro in real time. Respiratory waveforms were analyzed by the SASSI module of the Breathe Easy software to determine respiratory frequency (Vf), tidal volume (VT), minute ventilation (VE), oxygen consumption (VO2), ventilatory equivalents for oxygen (VE/VO2), inspiratory duration (TI), expiratory duration (TE), breath cycle duration (TTOT), inspiratory flow rate (VT/TI), and expiratory flow rate (VT/TE) (Lusk et al., 2023). Poincaré plots for TI, TE, TTOT, VT/TI, VT/TE, Vf, VT, VE and their SD1 and SD2 properties were generated and characterized by the STAGG module of the Breathe Easy software. Analysis windows of the respiratory parameters were as follows: for room air, the entire experimental period was considered; for hypercapnia, only the last 5 min challenge for each of 5/7/10% CO2 conditions were analyzed; and for hypoxia, three separate 5-min time periods, namely the 5th-10th min, 10th-15th min, and 15th-20th min intervals, were used due to the stereotypical biphasic respiratory response under this condition. Only steady quiescent breathing periods were included in the data analysis. A power analysis was performed using the reported effect size in Abbott et al. (2014), which used the same mouse model as we used here; 5-13 mice were necessary to observe a statistically significant result (Abbott et al. (2014) was able to see a significant difference between two groups with n=7 mice). In our experiments, the sample size for each group (Dbh-Cre; Vglut2 cKO and control) exceeded 13. Steady state results (Vf, VE, VT, VO2, VE/VO2, TI, TE, TTOT, VT/TI, VT/TE) for room air and hypercapnic or hypoxic data were compared between Dbh-Cre; Vglut2 cKO cohorts and sibling controls using a linear mixed-effects regression model with animal type (experimental vs. control) as fixed effects and animal ID as a random effect (Lusk et al., 2023). SD1 and SD2 results for room air and hypercapnic or hypoxic data were compared between Dbh-Cre; Vglut2 cKO cohorts and sibling controls using a Mann-Whitney U test (Ferreira et al., 2022). A p-value threshold of p < 0.05 was used to test for statistical significance. Individual data points, means, and standard errors of the mean are shown on all charts. The graphs were plotted by Prism 8.

Image quantification

Images were taken by using a Zeiss upright epifluorescent microscope and a Zeiss LSM 880 with Airyscan FAST confocal microscope. Images were captured using Zen software with z-stack function from top to bottom with 0.34μm intervals, exported, and then analyzed in Imaris using the spots and surface functions. For quantification of immunofluorescence staining, each GFP positive area coincident with DAPI (to denote nuclear localization) was defined as Vglut2 or Vglut3 expressing NA neurons while each tdTomato positive area coincident with DAPI was defined as NA neurons without any Vglut2 or Vglut3 co-expression. For quantification of Slc17a6 (Vglut2) or Slc17a8 (Vglut3) with Dbh double in situ hybridization in adult WT mice, Dbh positive areas coinciding with DAPI were identified as NA neurons and the Dbh positive areas overlapped with Slc17a6 (Vglut2) or Slc17a8 (Vglut3) positive pixels and DAPI were defined as NA neurons colocalized with Slc17a6 (Vglut2) or Slc17a8 (Vglut3). The number of Slc17a6 (Vglut2) or Slc17a8 (Vglut3) positive NA neurons and Slc17a6 (Vglut2) or Slc17a8 (Vglut3) negative NA neurons was counted in each image for both immunofluorescence and in situ experiments and the percentage of Slc17a6 (Vglut2) or Slc17a8 (Vglut3) positive NA neurons among all NA neurons in each NA nucleus in each mouse brainstem was calculated every other brain section unilaterally. For quantification of Slc17a6 (Vglut2) and Dbh double ISH in Dbh-Cre; Vglut2 cKO and their littermate controls, Slc17a6 (Vglut2) pixel intensities in Dbh positive areas coincident with DAPI (NA neurons) were measured in mutant and control images separately. The Slc17a6 (Vglut2) pixel intensity in NA neurons of control brains was normalized as 1 and the relative Slc17a6 (Vglut2) pixel intensity of NA neurons in mutant brains compared to that in control brains was calculated. At least three mouse brains were examined for each set of experiments in each group. All quantitative results were graphed using Prism 8.

Added in the revised version: Data availability

Added in the revised version: All the data supporting this study is available from the corresponding author upon request.